# These tests are meant to go hand-in-hand with the rendered alert rule.
# Use github.com/scality/action-prom-render-test@python-renderer python module
#
# Render the alerts file with
#   alertgen alerts.yaml -i 'namespace=zenko,service=data-db-mongodb-sharded,pvc=datadir-mongodb,replicas=3'

evaluation_interval: 30s

rule_files:
  - alerts.rendered.yaml

tests:
  - name: MongoDbDown
    interval: 30s
    input_series:
      - series: up{namespace="zenko", job="zenko/data-db-mongodb-sharded-mongos", pod="data-db-mongodb-sharded-mongos-0"}
        values: 1 1 0 0 0 0 1 1 1
      - series: up{namespace="zenko", job="zenko/data-db-mongodb-sharded-mongos", pod="data-db-mongodb-sharded-mongos-1"}
        values: 1 1 1 1 0 0 0 1 1
      - series: up{namespace="zenko", job="zenko/data-db-mongodb-sharded-mongos", pod="data-db-mongodb-sharded-mongos-2"}
        values: 1 1 1 1 1 0 0 0 1
    alert_rule_test:
      - { alertname: MongoDbDegraded, eval_time: 30s, exp_alerts: [] }
      - { alertname: MongoDbCritical, eval_time: 30s, exp_alerts: [] }

      - { alertname: MongoDbDegraded, eval_time: 60s, exp_alerts: [] }
      - { alertname: MongoDbCritical, eval_time: 60s, exp_alerts: [] }

      - { alertname: MongoDbCritical, eval_time: 90s, exp_alerts: [] }
      - alertname: MongoDbDegraded
        eval_time: 90s
        exp_alerts:
          - exp_labels:
              severity: warning
              job: mongos
            exp_annotations:
              description: "Less than 100% of MongoDb mongos instances are up and healthy: 2/3."
              summary: MongoDb mongos service degraded

      - { alertname: MongoDbCritical, eval_time: 120s, exp_alerts: [] }
      - alertname: MongoDbDegraded
        eval_time: 120s
        exp_alerts:
          - exp_labels:
              severity: warning
              job: mongos
            exp_annotations:
              description: "Less than 100% of MongoDb mongos instances are up and healthy: 1/3."
              summary: MongoDb mongos service degraded

      - alertname: MongoDbDegraded
        eval_time: 150s
        exp_alerts:
          - exp_labels:
              severity: warning
              job: mongos
            exp_annotations:
              description: "Less than 100% of MongoDb mongos instances are up and healthy: 0/3."
              summary: MongoDb mongos service degraded
      - alertname: MongoDbCritical
        eval_time: 150s
        exp_alerts:
          - exp_labels:
              severity: critical
              job: mongos
            exp_annotations:
              description: "Less than 50% of MongoDb mongos instances are up and healthy: 0/3."
              summary: MongoDb mongos service critical
      
      - alertname: MongoDbDegraded
        eval_time: 180s
        exp_alerts:
          - exp_labels:
              severity: warning
              job: mongos
            exp_annotations:
              description: "Less than 100% of MongoDb mongos instances are up and healthy: 1/3."
              summary: MongoDb mongos service degraded
      - alertname: MongoDbCritical
        eval_time: 180s
        exp_alerts:
          - exp_labels:
              severity: critical
              job: mongos
            exp_annotations:
              description: "Less than 50% of MongoDb mongos instances are up and healthy: 1/3."
              summary: MongoDb mongos service critical

  - name: NoPrimary
    interval: 1m
    input_series:
      - series: mongodb_rs_members_state{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-0",member_state="PRIMARY"}
        values: 1 _

    alert_rule_test:
      - alertname: NoPrimary
        eval_time: 1m
        exp_alerts: []
      - alertname: NoPrimary
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              namespace: zenko
              member_state: PRIMARY
              severity: critical
            exp_annotations:
              description: "MongoDb has no PRIMARY."
              summary: MongoDb down

  - name: UnhealthyMemberWarning
    interval: 30s
    input_series:
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-0", member_idx="data-db-mongodb-sharded-shardsrv-0.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 1 1 1 1
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-0", member_idx="data-db-mongodb-sharded-shardsrv-1.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 1 1
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-0", member_idx="data-db-mongodb-sharded-shardsrv-2.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 0 0
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-1", member_idx="data-db-mongodb-sharded-shardsrv-0.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 1 1 1 1
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-1", member_idx="data-db-mongodb-sharded-shardsrv-1.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 1 1
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-1", member_idx="data-db-mongodb-sharded-shardsrv-2.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 0 0
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-2", member_idx="data-db-mongodb-sharded-shardsrv-0.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 1 1 _ _
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-2", member_idx="data-db-mongodb-sharded-shardsrv-1.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 _ _
      - series: mongodb_rs_members_health{namespace="zenko",pod="data-db-mongodb-sharded-shardsrv-2", member_idx="data-db-mongodb-sharded-shardsrv-2.data-db-mongodb-sharded-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 1 1 _ _

    alert_rule_test:
      - alertname: UnhealthyMemberWarning
        eval_time: 1m
        exp_alerts: []
      - alertname: UnhealthyMemberWarning
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              member_idx: data-db-mongodb-sharded-shardsrv-2
              severity: warning
              member_state: SECONDARY
            exp_annotations:
              description: Member data-db-mongodb-sharded-shardsrv-2 (SECONDARY) is not healthy.
              summary: Unhealthy MongoDb member

  - interval: 30s
    input_series:
      - series: mongodb_rs_members_electionDate{namespace="zenko", pod="data-db-mongodb-sharded-shard0-data-0", job="zenko/data-db-mongodb-sharded-shard0-data"}
        values: 1+0x9 10+1x10
      - series: mongodb_rs_members_electionDate{namespace="zenko", pod="data-db-mongodb-sharded-cfgsvr-0", job="zenko/data-db-mongodb-sharded-cfgsvr"}
        values: 1+0x20

    alert_rule_test:
      - alertname: TooManyElectionsWarning
        eval_time: 9m
        exp_alerts: []
      - alertname: TooManyElectionsWarning
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              job: shard0-data
            exp_annotations:
              description: Number of elections in last 10m is greater than 10 for `shard0-data`.
              summary: Too many elections

  - name: ReplicationLagWarning
    interval: 1m
    input_series:
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-0", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-0.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 5  35000+1000x10 45000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-0", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-1.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 0  24000+1000x10 39000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-0", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-2.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 2   2000+1000x10 41000

      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-1", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-0.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 5  34000+1000x10 44000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-1", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-1.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 0  26000+1000x10 40000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-1", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-2.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 2   1000+1000x10 40000

      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-2", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-0.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 5  12000+1000x10 43000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-2", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-1.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 0   9000+1000x10 38000
      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-shard0-data-2", rs_nm="shard0-data", member_idx="data-db-mongodb-sharded-shard0-data-2.data-db-mongodb-sharded-shardsrv-headless.svc.cluster.local", member_state="SECONDARY"}
        values: 2   3000+1000x10 42000

      - series: mongodb_rs_members_optimeDate{namespace="zenko",pod="data-db-mongodb-sharded-configsvr-1", rs_nm="configsvr", member_idx="data-db-mongodb-sharded-cfgsvr-0.data-db-mongodb-sharded-cfgsvr-headless.svc.cluster.local", member_state="PRIMARY"}
        values: 71  8300  9500

    alert_rule_test:
      - alertname: ReplicationLagWarning
        eval_time: 1m
        exp_alerts: []
      - alertname: ReplicationLagWarning
        eval_time: 2m
        exp_alerts: []
      - alertname: ReplicationLagWarning
        eval_time: 10m
        exp_alerts: []
      - alertname: ReplicationLagWarning
        eval_time: 11m
        exp_alerts:
          - exp_labels:
              severity: warning
              member_idx: data-db-mongodb-sharded-shard0-data-2
              rs_nm: shard0-data
            exp_annotations:
              description: Mongodb replication lag for `data-db-mongodb-sharded-shard0-data-2` is more than 30 seconds.
              summary: MongoDB replication lag
      - alertname: ReplicationLagWarning
        eval_time: 12m
        exp_alerts: []

  - name: TooManyClientConnectionsWarning
    interval: 1m
    input_series:
      - series: mongodb_ss_connections{namespace="zenko",pod="data-db-mongodb-sharded-mongos-0", job="zenko/data-db-mongodb-sharded-mongos", conn_type="current"}
        values: 90 98 110 110

    alert_rule_test:
      - alertname: TooManyClientConnectionsWarning
        eval_time: 2m
        exp_alerts: []
      - alertname: TooManyClientConnectionsWarning
        eval_time: 4m
        exp_alerts:
          - exp_labels:
              severity: warning
              job: mongos
            exp_annotations:
              description: Too many client connections on `mongos`.
              summary: Too many MongoDB client connections

  - name: RemainingDiskSpaceWarning
    interval: 1m
    input_series:
      - series: kubelet_volume_stats_available_bytes{namespace="zenko",persistentvolumeclaim="datadir-mongodb-0"}
        values: 30-5x3
      - series: kubelet_volume_stats_capacity_bytes{namespace="zenko",persistentvolumeclaim="datadir-mongodb-0"}
        values: 100+0x3

    alert_rule_test:
      - alertname: RemainingDiskSpaceWarning
        eval_time: 2m
        exp_alerts: []
      - alertname: RemainingDiskSpaceWarning
        eval_time: 4m
        exp_alerts:
          - exp_labels:
              severity: warning
              namespace: zenko
              persistentvolumeclaim: datadir-mongodb-0
            exp_annotations:
              description: MongoDb has low disk space
              summary: MongoDb has low disk space
    

  - name: MongoDbPodRecovering
    interval: 1m
    input_series:
      - series: mongodb_rs_members_state{namespace="zenko", pod="data-db-mongodb-sharded-mongos-0", rs_state="3"}
        values: 3x1440

    alert_rule_test:
      - alertname: MongoDbPodRecovering
        eval_time: 1d
        exp_alerts:
          - exp_labels:
              namespace: zenko
              severity: critical
              pod: data-db-mongodb-sharded-mongos-0
              rs_state: 3
            exp_annotations:
              description: "MongoDB pod `data-db-mongodb-sharded-mongos-0` has been in the 'RECOVERING' state more than 24 hours. The instance may be failing to catch up and recover."
              summary: MongoDB is recovering
          - exp_labels:
              namespace: zenko
              pod: data-db-mongodb-sharded-mongos-0
              rs_state: 3
              severity: warning
            exp_annotations:
              description: "MongoDB pod `data-db-mongodb-sharded-mongos-0` has been in 'RECOVERING' state for more than 1 hour. This may be expected if the 'Resync a Data Services MongoDB Member' procedure has recently been executed."
              summary: MongoDB is recovering

  - name: MongoDbInvalidState
    interval: 1m
    input_series:
      - series: mongodb_rs_members_state{namespace="zenko", pod="data-db-mongodb-sharded-mongos-0", rs_state="6"}
        values: 6x10
      - series: mongodb_rs_members_state{namespace="zenko", pod="data-db-mongodb-sharded-mongos-0", rs_state="8"}
        values: 0x10
      - series: mongodb_rs_members_state{namespace="zenko", pod="data-db-mongodb-sharded-mongos-0", rs_state="10"}
        values: 0x10

    alert_rule_test:
      - alertname: MongoDbInvalidState
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              namespace: zenko
              severity: critical
              pod: data-db-mongodb-sharded-mongos-0
              rs_state: 6
            exp_annotations:
              description: "MongoDB pod `data-db-mongodb-sharded-mongos-0` is in an invalid state state (`6`)."
              summary: "MongoDB node in an invalid state: 6 (UNKNOWN), 8 (DOWN), 10 (REMOVED)"

  - name: MongoDbPodStartup2
    interval: 1m
    input_series:
      - series: mongodb_rs_members_state{namespace="zenko", pod="data-db-mongodb-sharded-mongos-0", rs_state="5"}
        values: 5x60

    alert_rule_test:
      - alertname: MongoDbPodStartup2
        eval_time: 1h
        exp_alerts:
          - exp_labels:
              namespace: zenko
              severity: warning
              pod: data-db-mongodb-sharded-mongos-0
              rs_state: 5
            exp_annotations:
              description: "MongoDB pod `data-db-mongodb-sharded-mongos-0` has been in the 'STARTUP2' state for more than 1 hour. Please ensure that the instance is running properly."
              summary: MongoDB node in STARTUP2 state for too long
    

  - name: MongoDbRSNotSynced
    interval: 1m
    input_series:
      - series: mongodb_rs_members_state{namespace="zenko", rs_nm="rs0", pod="data-db-mongodb-sharded-mongos-2", member_state="SECONDARY"}
        values: 1x10
  
    alert_rule_test:
      - alertname: MongoDbRSNotSynced
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              rs_nm: rs0
            exp_annotations:
              description: "MongoDB replica set `rs0` is not in the expected state. It currently has `1` SECONDARY members instead of the expected number. Please ensure that all instance are running properly."
              summary: MongoDB replica set out of sync



